# Loops

Topics:
1. `while` loops
2. `for` loops
<br>

## while loops

The syntax for a `while` is

```julia
while *condition*
    *loop body*
end
```

For example, we could use `while` to count or to iterate over an array.

```julia
n = 0
while n < 10
    n += 1
    println(n)
end
n
```

Output:
```
1
2
3
4
5
6
7
8
9
10
10
```

```julia
fib = [1,2]

while length(fib) < 10
    push!(fib, fib[end]+fib[end-1])
    println(fib)

end
```

Output:
```
[1, 2, 3]
[1, 2, 3, 5]
[1, 2, 3, 5, 8]
[1, 2, 3, 5, 8, 13]
[1, 2, 3, 5, 8, 13, 21]
[1, 2, 3, 5, 8, 13, 21, 34]
[1, 2, 3, 5, 8, 13, 21, 34, 55]
[1, 2, 3, 5, 8, 13, 21, 34, 55, 89]
```

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]

i = 1
while i <= length(myfriends)
    friend = myfriends[i]
    println("Hi $friend, it's great to see you!")
    i += 1
end
```

Output:
```
Hi Ted, it's great to see you!
Hi Robyn, it's great to see you!
Hi Barney, it's great to see you!
Hi Lily, it's great to see you!
Hi Marshall, it's great to see you!
```

## for loops

The syntax for a `for` loop is

```julia
for *var* in *loop iterable*
    *loop body*
end
```

We could use a for loop to generate the same results as either of the examples above:

```julia
for n in 1:10
    println(n)
end
```

Output:
```
1
2
3
4
5
6
7
8
9
10
```

```julia
fib = [0,1]
for n in 3:10
    push!(fib, fib[end] + fib[end-1])
    println(fib)
end
```

Output:
```
[0, 1, 1]
[0, 1, 1, 2]
[0, 1, 1, 2, 3]
[0, 1, 1, 2, 3, 5]
[0, 1, 1, 2, 3, 5, 8]
[0, 1, 1, 2, 3, 5, 8, 13]
[0, 1, 1, 2, 3, 5, 8, 13, 21]
[0, 1, 1, 2, 3, 5, 8, 13, 21, 34]
```

```julia
for i in fib
    println(i/2)
end
```

Output:
```
0.0
0.5
0.5
1.0
1.5
2.5
4.0
6.5
10.5
17.0
```

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]

for friend in myfriends
    println("Hi $friend, it's great to see you!")
end
```

Output:
```
Hi Ted, it's great to see you!
Hi Robyn, it's great to see you!
Hi Barney, it's great to see you!
Hi Lily, it's great to see you!
Hi Marshall, it's great to see you!
```

Now let's use `for` loops to create some addition tables, where the value of every entry is the sum of its row and column indices. <br>

Note that we iterate over this array via column-major loops in order to get the best performance. More information about fast indexing of multidimensional arrays inside nested loops can be found at https://docs.julialang.org/en/v1/manual/performance-tips/#Access-arrays-in-memory-order,-along-columns-1

First, we initialize an array with zeros.

```julia
m, n = 5, 5
A = fill(0, (m, n))
```

Output:
```
5×5 Array{Int64,2}:
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
```

```julia
@time(
for j in 1:n
    for i in 1:m
        A[i, j] = i + j
    end
end)
A
```

Output:
```
0.000015 seconds (36 allocations: 1.125 KiB)
5×5 Array{Int64,2}:
 2  3  4  5   6
 3  4  5  6   7
 4  5  6  7   8
 5  6  7  8   9
 6  7  8  9  10
```

Here's some syntactic sugar for the same nested `for` loop

```julia
B = fill(0, (m, n))
```

Output:
```
5×5 Array{Int64,2}:
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
```

```julia
@time(for j in 1:n, i in 1:m
    B[i, j] = i + j
        end)
B
```

Output:
```
0.000015 seconds (36 allocations: 1.125 KiB)
5×5 Array{Int64,2}:
 2  3  4  5   6
 3  4  5  6   7
 4  5  6  7   8
 5  6  7  8   9
 6  7  8  9  10
```

The more "Julia" way to create this addition table would have been with an *array comprehension*.

```julia
C = [i + j for i in 1:m, j in 1:n]
```

Output:
```
5×5 Array{Int64,2}:
 2  3  4  5   6
 3  4  5  6   7
 4  5  6  7   8
 5  6  7  8   9
 6  7  8  9  10
```

### Exercises

#### 4.1 
Loop over integers between 1 and 100 and print their squares.

```julia
println([i^2 for i in 1:100])
```

Output:
```
[1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121, 144, 169, 196, 225, 256, 289, 324, 361, 400, 441, 484, 529, 576, 625, 676, 729, 784, 841, 900, 961, 1024, 1089, 1156, 1225, 1296, 1369, 1444, 1521, 1600, 1681, 1764, 1849, 1936, 2025, 2116, 2209, 2304, 2401, 2500, 2601, 2704, 2809, 2916, 3025, 3136, 3249, 3364, 3481, 3600, 3721, 3844, 3969, 4096, 4225, 4356, 4489, 4624, 4761, 4900, 5041, 5184, 5329, 5476, 5625, 5776, 5929, 6084, 6241, 6400, 6561, 6724, 6889, 7056, 7225, 7396, 7569, 7744, 7921, 8100, 8281, 8464, 8649, 8836, 9025, 9216, 9409, 9604, 9801, 10000]
```

#### 4.2 
Add to the code above a bit to create a dictionary, `squares` that holds integers and their squares as key, value pairs such that

```julia
squares[10] == 100
```

```julia
squares = Dict([(i, i^2) for i in 1:11])
```

Output:
```
Dict{Int64,Int64} with 11 entries:
  2  => 4
  11 => 121
  7  => 49
  9  => 81
  10 => 100
  8  => 64
  6  => 36
  4  => 16
  3  => 9
  5  => 25
  1  => 1
```

<!-- grade_id: cell-54834a2e4e122b70, points: 1, grade: true, solution: false -->
```julia
@assert squares[10] == 100
@assert squares[11] == 121
```

#### 4.3 
Use an array comprehension to create an an array `squares_arr` that stores the squares for all integers between 1 and 100.

```julia
squares_arr = [i^2 for i in 1:100]
```

Output:
```
100-element Array{Int64,1}:
     1
     4
     9
    16
    25
    36
    49
    64
    81
   100
   121
   144
   169
     ⋮
  7921
  8100
  8281
  8464
  8649
  8836
  9025
  9216
  9409
  9604
  9801
 10000
```

<!-- grade_id: cell-17d4631c51b8b4f7, points: 1, grade: true, solution: false -->
```julia
@assert length(squares_arr) == 100
@assert sum(squares_arr) == 338350
```

Please click on `Validate` on the top, once you are done with the exercises.